### Grading
The final score that you will receive for your programming assignment is generated in relation to the total points set in your programming assignment item—not the total point value in the nbgrader notebook.<br>
When calculating the final score shown to learners, the programming assignment takes the percentage of earned points vs. the total points provided by nbgrader and returns a score matching the equivalent percentage of the point value for the programming assignment. <br>
**DO NOT CHANGE VARIABLE OR METHOD SIGNATURES** The autograder will not work properly if your change the variable or method signatures. 

### WARNING
Please refrain from using **print statements/anything that dumps large outputs(>1000 lines) to STDOUT** to avoid running to into **memory issues**. 
Doing so requires your entire lab to be reset which may also result in loss of progress and you will be required to reach out to Coursera for assistance with this.
This process usually takes time causing delays to your submission.

### Validate Button
Please note that this assignment uses nbgrader to facilitate grading. You will see a **validate button** at the top of your Jupyter notebook. If you hit this button, it will run tests cases for the lab that aren't hidden. It is good to use the validate button before submitting the lab. Do know that the labs in the course contain hidden test cases. The validate button will not let you know whether these test cases pass. After submitting your lab, you can see more information about these hidden test cases in the Grader Output. <br>
***Cells with longer execution times will cause the validate button to time out and freeze. Please know that if you run into Validate time-outs, it will not affect the final submission grading.*** <br>

# EDA, Simple Linear Regression

In this assignment, we will use a simplified data and create a simple linear regression model. The dataset can be downloaded from https://www.kaggle.com/harlfoxem/housesalesprediction/download.    
This dataset contains house sale prices for Kings County, which includes Seattle. It includes homes sold between May 2014 and May 2015. There are several versions of the data. Some additional information about the columns is available here: https://geodacenter.github.io/data-and-lab/KingCounty-HouseSales2015/, some of which are copied below.

|Variable |	Description|
|:---------|:-------------|
|id 	|Identification|
|date |	Date sold|
|price |	Sale price|
|bedrooms |	Number of bedrooms|
|bathrooms |	Number of bathrooms|
|sqft_liv |	Size of living area in square feet|
|sqft_lot| 	Size of the lot in square feet|
|floors |	Number of floors|
|waterfront |	‘1’ if the property has a waterfront, ‘0’ if not.|
|view |	An index from 0 to 4 of how good the view of the property was|
|condition |	Condition of the house, ranked from 1 to 5|
|grade |	Classification by construction quality which refers to the types of materials used and the quality of workmanship. Buildings of better quality (higher grade) cost more to build per unit of measure and command higher value.|
|sqft_above |	Square feet above ground|
|sqft_basmt |	Square feet below ground|
|yr_built 	|Year built|
|yr_renov |	Year renovated. ‘0’ if never renovated|
|zipcode |	5 digit zip code|
|lat 	|Latitude|
|long 	|Longitude|
|squft_liv15 |	Average size of interior housing living space for the closest 15 houses, in square feet|
|squft_lot15 |	Average size of land lost for the closest 15 houses, in square feet|

```python
import scipy as sp
import scipy.stats as stats
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns
import copy
# Set color map to have light blue background
sns.set()
import statsmodels.formula.api as smf
import statsmodels.api as sm
%matplotlib inline
```

## 1. Munging data [15 pts]
In this part, let's load and inspect data. We will also learn how to transform columns when needed.

> 💡***Tip:***   `pd.read_csv(<file path>)` reads a csv file and returns to pandas data frame object. It can also read files with other delimiter such as `.tsv` files. pandas also has `pd.read_excel` to read excel files.
[See more in the document](https://pandas.pydata.org/docs/reference/api/pandas.read_csv.html)

```python
df = pd.read_csv('data/house_data_washington.csv')
```

### 1a) Date string to numbers [5 pts]
Now, let's overview the dataframe. Using `.head()` on the dataframe, we can see the first 5 rows of the data. You can specify number of rows as argument then it will show those number of rows. similarly, `.tail()` gives the last 5 rows by default. You can see the columns names, but not all columns are displayed if there are too many columns.  

> 💡***Tip:*** If you want to show all columns and rows, there are [pandas command](https://www.geeksforgeeks.org/show-all-columns-of-pandas-dataframe-in-jupyter-notebook/#) setting max rows and cols. Please do not submit your homework notebook with displaying large dataframe because it may crash from large memory consumption.  

The column 'date' is the date sold (with some black timestamp as well), and the data is string type (Note that sometimes data tables may have date/time columns as datetime object types. In this example data, it has a string type). We will extract year and month information from the string. 
In the data frame `df`, let's create new features 'sales_year' and 'sales_month' using 'date' column.

> 💡***Tip:***  You can use either bracket (e.g. `df['date']`) or dot (e.g. `df.date`) to get the column `'date'` in the data frame `df`. A single columns object from dataframe is a pandas series object type, and you can use `.apply()` method for a transformation. `.apply()` is generic and can be applied to not only to single column (pandas series) but also to multiple columns (pandas dataframe). Here, we will apply it to a single column object and use `lambda` function inside the .`apply()` as shown below. For more examples, see the [doc](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.apply.html).  

In this case, when we inspect the 'date' column, it is a string object, so we can slice the year and month from the string. Also, we'd like to convert the extracted year and month strings to ingeters.

```python
df.head()
```

```python
print(df.date)
print(type(df.date.iloc[0]))
```

```python
# extract year and month info from the string
# create new features 'sales_year' and 'sales_month' in df

df['sales_year'] = df.date.apply(lambda x: int(x[:4]))
df['sales_month'] = df.date.apply(lambda x: int(x[4:6]))
```

```python
df.groupby('sales_month')
```

Now, let's count how many sales occurred in each month and each year. We can use `.groupby()` function to group by 'sales_month' and 'sales_year' as shown below. 

> 💡***Tip:*** `.groupby()` itself returns an object that doesn't get displayed, not a processed dataframe. It is often used with some other aggregation method, such as `.count()`, `.mean()`, `.sum()`, etc. In the below example use case, we use `.count()` to count number of sales per each group (e.g. by sales_month). `.groupby()` can also group by multiple columns. This [resource](https://realpython.com/pandas-groupby/) has more explanations and examples.

```python
print(df.groupby('sales_month')['id'].count())
print(df.groupby('sales_year')['id'].count())
```

**Question 1a-1.** Based on the output from above cell, which month has the most number of sales?

<!-- grade_id: cell-447e548e063eb680, points: 0, grade: false, solution: true -->
```python
# your code here

# uncomment below and update the value as an integer
# most_sales_month = 
```

<!-- grade_id: cell-cadc8a85fd0b5630, points: 2.5, grade: true, solution: false -->
```python
# tests solutions for most_sales 
```

**Question 1a-2.** Which months has the least number of sales?

<!-- grade_id: cell-05ff2ec904565f4c, points: 0, grade: false, solution: true -->
```python
# your code here

# uncomment below and update the value as an integer
# least_sales_month = 
```

<!-- grade_id: cell-4b8eb3fefb462f4a, points: 2.5, grade: true, solution: false -->
```python
# tests solutions for least_sales
```

Now, let's have a look at what data type each columns has. We can use `.info()` method on the dataframe object to see the data type. You can see `int64`, `float64` and `object` in our example. `object` can be string type or something else (such as list or other types of objects). 
> 💡***Tip:*** Note that sometimes raw data is not adequately formatted that you might see columns that are supposed to be numbers can be typed as strings. It is a good practice to inspect data columns's data types and clean them if necessary.

```python
df.info()
```

### 1b) Variable types [5 pts]
Inspect each feature's data type and variable type. What is the best description for the variable type of following features? Update the string to 'numeric' or 'categorical'.

**Self-check Concepts**    
✔️ What data types can be considered as a numeric variable?    
✔️ What is the difference between ordinal and non-ordinal categorical variables?     

> 💡***Tip:*** Is binary categorical variable (Yes/No, Male/Female, True/False, Positive/Negative etc) numeric? Why or why not?      
How about a variable that has meaning of degree- such as survey/review ratings (very satisfied = 5, satisfied = 4, neutral = 3, disatisfied = 2, very disatisfied = 1)?         
Typically it is recommended to treat ordinal categorical variable (which order has meaning- e.g. degree, grades, numbers, severity etc) as numeric variable because a linear regression (or any ML) model can treat that variable (feature) as numbers and can learn a relationship to the target variable y. Also, categorical variables need to be binarized (which involves to transform the column into multiple binary columns) before used in a linear regression model. So if we treat an ordinal categorical variable to one numeric variable column instead of multiple binary columns representing categorical variable, it is more efficient for the model. Remember- a simpler model with the same information is better! 

<!-- grade_id: cell-d1457a08b8feb7e4, points: 0, grade: false, solution: true -->
```python
# your code here


# uncomment the feaures below and update the strings with 'numeric' or 'categorical'
# price = ''
# bathrooms = ''
# waterfront = ''
# grade = ''
# zipcode = ''
# sales_year = ''
```

```python
# You can use below to check what unique values exist in each column.
for c in df.columns[2:]:
    print(c, df[c].unique())

```

<!-- grade_id: cell-83405385a8bf3251, points: 0, grade: true, solution: false -->
```python
# test 1/6 for 1b 
# This test is an example. The rest tests are hidden but are the same as this one checking each answer.
assert price == 'numeric', "Check 1b. What is the correct variable type for price?"
```

<!-- grade_id: cell-85adebcef83311e5, points: 1, grade: true, solution: false -->
```python
# test 2/6 for 1b 
```

<!-- grade_id: cell-1a72aecc2ae2a96c, points: 1, grade: true, solution: false -->
```python
# test 3/6 for 1b 
```

<!-- grade_id: cell-0a727f18811562c2, points: 1, grade: true, solution: false -->
```python
# test 4/6 for 1b 
```

<!-- grade_id: cell-bb80917a7ee3c5b7, points: 1, grade: true, solution: false -->
```python
# test 5/6 for 1b 
```

<!-- grade_id: cell-c35914383b2fba6f, points: 1, grade: true, solution: false -->
```python
# test 6/6 for 1b 
```

### 1c) Drop features [5 pts]
Let's drop features that are unnecessary. `id` is not a meaningful feature. `date` string has been coded to `sales_month` and `sales_year`, so we can remove `date`. `zipcode` can be also removed as it's hard to include in a linear regression model and the location info is included in the `lat` and `long`.
Drop the features `id`, `date`, and `zipcode` and replace the df.

> 💡***Tip:*** .drop() function can drop one or more columns or rows. Learn how to use it in the [doc](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.drop.html)


<!-- grade_id: cell-fba7fcd19732194e, points: 0, grade: false, solution: true -->
```python
# drop unnecessary features, replace df
# your code here


```

<!-- grade_id: cell-09acbfd67a77b038, points: 5, grade: true, solution: false -->
```python
# tests that you droppd the features id, date, and zipcode from df
```

## 2. More inspection; Correlation and pair plot [5 pts]
### 2a) Get correlation matrix on the data frame. [5 pts]
Which feature may be the best predictor of price based on the correlation? Answer as a string value (e.g. best_guess_predictor = 'price' or best_guess_predictor = 'yr_built')

> 💡***Tip:*** .corr() finction can show correlation matrix from the dataframe. [More resource](https://www.geeksforgeeks.org/python-pandas-dataframe-corr/#)


**Self-check Concepts**    
✔️ By looking at the correlation matrix, how do you decide which feature is the best predictor?

```python
df.corr()
```

<!-- grade_id: cell-0e9bdc9ee6161b7e, points: 0, grade: false, solution: true -->
```python
# your code here

# uncomment and update best_guess_predictor with a string value
# best_guess_predictor = ''
```

<!-- grade_id: cell-455e59f31501b283, points: 5, grade: true, solution: false -->
```python
# tests the solution for best_guess_predictor
```

### 2b) Display the correlation matrix as heat map [Not graded]
[`seaborn.heatmap()`](https://seaborn.pydata.org/generated/seaborn.heatmap.html) can visualize a matrix as a heatmap. Visualize the correlation matrix using seaborn.heatmap(). Play with color map, text font size, decimals, text orientation etc. For example, the resulting display may look like below. If you find how to make a pretty visualization, please share in the discussion board.
![correlation matrix](imgs/correlation_matrix.png)
<br>
**Note:** your code for this section may cause the Validate button to time out. If you want to run the Validate button prior to submitting, you could comment out the code in this section.

<!-- grade_id: cell-2e19c30b0454fabd, points: 0, grade: false, solution: true -->
```python
# practice visualizing correlation matrix using a heatmap
# your code here

```

### 2c) Pair plot [Not graded]
Pair plot is a fast way to inspect relationships between features. Use seaborn's .pairplot() function to draw a pairplot if the first 10 columns (including price) and inspect their relationships. Set the diagonal elements to be KDE plot. The resulting plot will look like below.
![pair plot](./imgs/pair_plot.png)

**Note:** your code for this section may cause the Validate button to time out. If you want to run the Validate button prior to submitting, you could comment out the code in this section.

<!-- grade_id: cell-630895ed926e33b7, points: 0, grade: false, solution: true -->
```python
# practice inspecting relationships between features using a pair plot. 
# your code here

```

## 3. Simple linear regression [20 pts]

### 3a) Data preparation [5 pts]
We will split the data to train and test datasets such that the test dataset is 20% of original data.
Use `sklearn.model_selection.train_test_split` function to split the data frame to X_train and X_test. X_train is 80% of observation randomly chosen. X_test is the rest 20%. Both X_train and X_test are `pd.DataFrame` object and include 'price' in the table. Note that the train_test_split can handle data frame as well as array.

> 💡***Tip:*** Use sklearn.model_selecttion.train_test_split to split the data frame. We would like X_train to be 80% of the observation and X_test to be 20% of the observations. Print length of X_train and X_test.

<!-- grade_id: cell-039f3a4314195737, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split

# your code here



```

<!-- grade_id: cell-8e12fbb95b41c1fd, points: 1, grade: true, solution: false -->
```python
# Testing cell for self-check
assert(len(X_train) == 17290), "Check 3a, did you split properly so X_Train is 80% of the observations?"
assert(type(X_train)==type(pd.DataFrame())), "Check 3a, what type of object should X_train be?"
```

<!-- grade_id: cell-71c661acc821788e, points: 2, grade: true, solution: false -->
```python
# Testing cell 
```

<!-- grade_id: cell-ed5049ec43a9c89c, points: 2, grade: true, solution: false -->
```python
# Testing cell 

```

### 3b) Train a simple linear regression model [5 pts]
Use the best_guess_predictor as a single predictor and build a simple linear regression model using **`statsmodels.formula.api.ols`** function (https://www.statsmodels.org/dev/example_formulas.html)
Print out the result summary. Train on the X_train portion. What is the adjusted R-squared value?

> 💡***Tip:*** We had imported the library at the top of this notebook. So you can use the `smf` alias.
>```python
>import statsmodels.formula.api as smf
>```

N.B.: It recommended that you use the statsmodel library to do the regression analysis as opposed to e.g. sklearn. The sklearn library is great for advanced topics, but it's easier to get lost in a sea of details and it's not needed for these problems.

<!-- grade_id: cell-f5291d1c1d6ebee5, points: 0, grade: false, solution: true -->
```python
# use best_guess_predictor as a single predictor
# build a simple linear regression model, train on the X_train portion

# Make sure to use the `statsmodels.formula.api.ols` function for building the model. 
# model = 

#update following value according to the result
# adj_R2 =  

# your code here



```

<!-- grade_id: cell-2d149c760cefd97e, points: 1, grade: true, solution: false -->
```python
# self test
assert len(model.params.index) == 2, 'Check 3b, Number of model parameters (including intercept) does not match. Did you make a univariate model?'
```

<!-- grade_id: cell-fc6e4ae429f34164, points: 2, grade: true, solution: false -->
```python
# hidden test for 3b
```

<!-- grade_id: cell-5425e72049adbf2f, points: 2, grade: true, solution: false -->
```python
# hidden test for 3b
                                                                          
```

### 3c) Best predictor [10 pts]
In question 5a, we picked a best guess predictor for price based on the correlation matrix. Now we will consider whether the best_guess_predictor that we used is still the best.<br>
Print out a list ranking all of the predictors. Then print out a list of the top three predictors in order.<br>      
**Hint:** Linear regression uses adjusted R squared as fit performance. So you can rank by this metric.<br>
- What were your top three predictors? 
- How did you order your list of predictors to select those as the top ones? 
- Is your top predictor for this section the same as the best guess predictor you selected in question 2a? 

<!-- grade_id: cell-9ca80c1b78b9ccd7, points: 0, grade: false, solution: true -->
```python
# your code here

# uncomment and update top_three
# top_three = []
```

<!-- grade_id: cell-385027dd29c1bbd3, points: 1, grade: true, solution: false -->
```python
# self test cell
assert(type(top_three) == list), "Check 3c, the top_three needs to be a list."
assert(len(top_three) == 3), "Check 3c, the top_three list needs to have three element."
```

<!-- grade_id: cell-823971647e6e28eb, points: 2, grade: true, solution: false -->
```python
# test cell
```

<!-- grade_id: cell-d817fe4469b808c2, points: 2, grade: true, solution: false -->
```python
# test cell
```

<!-- grade_id: cell-4250799ed572919a, points: 2, grade: true, solution: false -->
```python
# test cell
```

<!-- grade_id: cell-a1c9231f7c678cca, points: 1, grade: true, solution: false -->
```python
# test cell
```

<!-- grade_id: cell-c2d5eb3a0be8fe42, points: 1, grade: true, solution: false -->
```python
# test cell
```

<!-- grade_id: cell-5067d7fbced43b91, points: 1, grade: true, solution: false -->
```python
# test cell
```

```python

```