# Probabilistic Language Modeling

## Conditional Probability

$P(A \mid B)=\frac{P(A \cap B)}{P(B)}$

$P(A \cap B)=P(B \mid A) * P(A)$

```python
from collections import Counter
from typing import List, Union, Tuple
```

<!-- grade_id: cell-b028f34ad0c3f8f5, points: 0, grade: false, solution: true -->
```python
def prob(word: str, corpus: List[str]) -> float:
    """Return the probability of `word` in `corpus`.
    P(w_1) = count(w_1) / N
    """
    
    # Below gives you a list of whitespace-separated tokens for all sentences
    # in the corpus, and a dictionary of counts for each unique word.
    tokens = [token for sentence in corpus for token in sentence.split()]
    counts = Counter(tokens)
    
    # Find the probability of word given these counts
    # your code here
#     raise NotImplementedError
    
    total_words = len(tokens)
    word_count = counts[word] if word in counts else 0
    probability = word_count / total_words if total_words > 0 else 0
    
    return probability

print(prob("this", ["this is a corpus"]))  # 0.25



```

Output:
```
0.25
```

<!-- grade_id: cell-ec7859188373645f, points: 5, grade: true, solution: false -->
```python
"""Test cases"""
assert prob("this", ["this is a corpus"]) == 0.25
print('All Test Cases Have Passed!')
```

Output:
```
All Test Cases Have Passed!
```

```python

```